## Emotion Detection and Recognition
Emotion Detection and Recognition from text is closely related to Sentiment Analysis which aims to detect and recognize various types of feelings through the expression of texts, such as anger, disgust, fear, happiness, sadness, and surprise unlike sentiment analysis which detects positive, neutral, or negative feelings from text.


## Data Preprocessing:

- Read the dataset ISEAR.csv which is in the current path.
- Visualise and Clean the dataset
- Perform any preprocessing that may be required

The Emotions in the dataset should be: 'joy', 'fear', 'anger', 'sadness', 'disgust', 'shame' and 'guilt'.

<!-- grade_id: cell-ee4c535e719b9073, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
import pandas as pd
df=pd.read_csv('ISEAR.csv',names=['SN','Emotions','Message'])
df.drop("SN", axis=1, inplace=True)
df.head()

```

Output:
```
Emotions                                            Message
0      joy  On days when I feel close to my partner and ot...
1     fear  Every time I imagine that someone I love or I ...
2    anger  When I had been obviously unjustly treated and...
3  sadness  When I think about the short time that we live...
4  disgust  At a gathering I found myself involuntarily si...
```

<!-- grade_id: cell-6ba89ba3bf85a0f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
Emotions=df['Emotions']
df.Emotions.unique()
df.shape
df.columns
df.isnull().sum()
Y=df.Emotions
X=df.Message

```

## Model Creation:

Create any model you wish that classifies emotions with an accuracy > 60%. 
You can use all the data in X in the training set itself or split it.
Save the model as model. Your model will be graded using model.predict on hidden dataset.
<br> Since the dataset is very less, Model with accuracy >60% is enough to pass the course.
<br> **We highly recommend you to test and use as many skills you have learned in this course and have explore further in the web.**

<!-- grade_id: cell-1dd10e27aeaaf5d2, points: 0, grade: false, solution: true -->
```python
model = None


#
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.feature_extraction.text import TfidfTransformer
from sklearn.pipeline import Pipeline
from sklearn.linear_model import SGDClassifier

model = Pipeline([('vect', CountVectorizer()),
                     ('tfidf', TfidfTransformer()),
                     ('clf', SGDClassifier(loss='hinge', penalty='l2',
                                           alpha=1e-3, max_iter=5, random_state=42)),
                     ])
model.fit(X, Y)


```

Output:
```
Pipeline(memory=None,
     steps=[('vect', CountVectorizer(analyzer='word', binary=False, decode_error='strict',
        dtype=<class 'numpy.int64'>, encoding='utf-8', input='content',
        lowercase=True, max_df=1.0, max_features=None, min_df=1,
        ngram_range=(1, 1), preprocessor=None, stop_words=None,
        strip...ty='l2', power_t=0.5, random_state=42, shuffle=True,
       tol=None, verbose=0, warm_start=False))])
```

<!-- grade_id: cell-4762880664504021, points: 10, grade: true, solution: false -->
```python
test_X =["I just won a Million Dollar. I don't know what to do with this money.",
        'After seeing a horror film - Omen.  I could not sleep at night even though I closed my eyes.  The last scene of the film kept coming back to me, I seemed to hear the girls in the film  \nshouting at the devil.  The devil did not have any expression on  \nhis face and I thought that he would choose me.']
test_y=['joy','fear']

prediction = model.predict(test_X)
assert(prediction[0]==test_y[0])
assert(prediction[1]==test_y[1])
```

Congratulations, you have reached the end of the Module Project